Negative affect influences the computations underlying food choice in bulimia nervosa

Author response:

The following is the authors’ response to the previous reviews

eLife Assessment

This study makes a valuable contribution to understanding how negative affect shapes food-choice decision making in bulimia nervosa by leveraging a mechanistic drift diffusion model to quantify the weighting of tastiness and healthiness attributes. The evidence is solid, supported by a randomized crossover design and generally appropriate statistical analyses. However, the interpretability of the findings is limited by ambiguities in the affect manipulation, particularly regarding whether neutral and negative inductions yielded reliably distinct affective states at the time of task performance in the bulimia nervosa group. Consequently, session-related differences in model parameters cannot be unequivocally attributed to negative affect rather than to uncontrolled state or contextual factors, and clearer separation of affective conditions alongside analyses aligned with the paired data structure would strengthen the conclusions.

We thank the Editor and Reviewers for their careful summary of the study's strengths and for their constructive feedback.

The eLife Assessment identified two specific limitations that qualified the strength of evidence:

(1) ambiguity regarding whether the two affect inductions yielded reliably distinct affective states in the BN group at the time of task performance, and (2) analyses that were not fully aligned with the paired data structure. We have directly addressed both concerns in this revision. We provide explicit statistical evidence confirming that neutral and negative inductions yielded distinct affective states in the bulimia nervosa group; and we have re-analyzed all DDM parameters using updated mixed-effects regressions with an unstructured covariance matrix that appropriately accounts for the paired data structure. For completeness, we have also added the requested difference-in-difference analysis. Both approaches yielded conclusions consistent with those originally reported.

In light of these revisions, we would be grateful if the Editorial Team would consider whether the strength of evidence rating might be updated from "solid" to "convincing." All changes in the revised manuscript are marked in blue.

Public Reviews:

Reviewer #1 (Public review):

Summary:

Using a computational modeling approach based on the Drift and Diffusion Model (DDM) introduced by Ratcliff and McKoon in 2008, the article by Shevlin and colleagues investigates whether there are differences between neutral and negative emotional states in:

(1) The timings of the integration in food choices of the perceived healthiness and tastiness of food options in individuals with bulimia nervosa (BN) and healthy participants (2) The weighting of the perceived healthiness and tastiness of these options.

Strengths:

By looking at the mechanistic part of the decision process, the approach has potential to improve the understanding of pathological food choices.

Weaknesses:

I thank the authors for revising their manuscript.

I still notice that the authors did not go through their manuscript to look for wordings refering to a prediction interpretation of their results while I already highlighted the inappropriateness of this wording in my two first rounds of reviews: e.g. there is still "we used zero-inflated negative binomial models to predict the three-month frequency" and I can find other statements like this. The design of their study does not allow such claims.

We thank the Reviewer for identifying cases where the term “predicted” may mislead readers about the causal nature of our claims. We have made the following edits (changes are italicized):

Methods (lines 516-518): “For these exploratory analyses, we used negative binomials to test the association between parameter estimates and the three-month frequency of retrospectively reported Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE).”

Figure 5 (lines 881-882): “Affect-induced changes in information onset were associated with more frequent subjective binge episodes.

The authors answered my major concern regarding the experimental induction towards a negative or a neutral state before running the food decision task. My concern is: BN patients already seemed to be already in a high negative state before undergoing the neutral induction, while these patients are in a lower negative state before undergoing the negative induction. It is therefore not surprising that patients seem to report a similar level of negative state after the two inductions (according to the figure of the authors' previous article). Of note is that the additional analysis the authors ran within the BN group only provides a significant result: this result shows that there has been an induction but does not rule out that patients were in the exact same magnitude of negative state to perform the task as the figure in their previously published article suggests it. The major issue is to show that:

(1) As compared to the neutral induction, there has been a higher variation in negative state after as compared to before the negative induction.

(2) The magnitude of the negative state after the negative induction is higher than the magnitude of the negative state after the neutral induction.

The first point shows that the induction worked. The second point shows that the participants are in two distinct states. Without showing the second point, it may be possible that one induction increases the negative state of participants to the same level as the one of the second induction that has not increased anything.

Within this context, how is it possible to associate, in patients, a difference in the DDM between the two sessions to a negative state (which is one of the main focus of the article) rather than to another parameter that has not been captured? A similar situation would be in an experiment studying the consequence of stress, a stressfull induction over relaxed participants attending the lab has high chances to raise the level of stress of those participants to the same level as the one that the same participants would experience after a neutral induction when these participants attend the lab with an already high level of stress. In that case, would it be approrpiate to claim that a difference at a task performed after the induction would be related to stress while the participants would be at the same level of stress when performing the task despite the fact that the induction worked ?

In the experiment performed by the authors, the additional analysis to perform would be a paired sample t-test (or the appropriate non-parametric test) to check whether the magnitude of negative state of BN patients was different between the negative and neutral conditions after the induction only. If not, associating the difference at the DDM with negative states in BN is highly misleading.

We thank the Reviewer for pressing on this point, and we apologize that our previous response did not make this sufficiently explicit. We agree with the Reviewer that two things must be demonstrated: (1) that the negative induction produced a greater change in negative affect than the neutral induction, and (2) that the magnitude of post-induction negative affect was higher following the negative induction than the neutral induction. We had included the results of analyses addressing both points in the Supplementary Materials of our previous submission, but we appreciate that we had not made this clear in our response.

Regarding point (1), the mixed-effects model in Supplementary Table S1 yielded a significant Affect Condition × Timing interaction (β = 20.43, SE = 6.35, t = 3.22, p = 0.002), confirming that negative affect increased significantly more from pre- to post-induction in the negative condition than in the neutral condition. This is further supported by within-BN-group analyses in the Supplementary Materials: the negative affect induction produced a large, significant increase in negative affect (mean difference = 20.36, SE = 4.21, t = 4.84, p < 0.0001, Cohen's d = 0.97), whereas the neutral induction was not associated with a significant change in negative affect (mean difference = 7.16, SE = 4.21, t = 1.70, p = 0.327, Cohen's d = 0.34).

Regarding point (2), we directly compared post-induction negative affect between conditions within the BN group, as requested by the Reviewer. The magnitude of negative affect was significantly higher following the negative mood induction than after the neutral mood induction (mean difference = 17.40, SE = 4.21, t = 4.13, p = 0.0003, Cohen's d = 0.83). This large effect size confirms that participants with BN were in meaningfully distinct affective states when performing the food decision task under the two conditions.

Together, these analyses establish (1) that the induction worked as intended, and (2) that the two post-induction states were both statistically and practically distinct. We have added explicit language to the manuscript to make both of these points clear (lines: 181-185):

Critically, post-induction negative affect within the BN group was significantly higher following the negative affect induction than after the neutral affect induction (mean difference = 17.40, SE = 4.21, t = 4.13, p < 0.001, Cohen's d = 0.83; see Supplementary Materials for full details), confirming that BN participants completed the food decision task under meaningfully distinct affective states across the two sessions.

I read carefully the authors' answer related to mixed models: they claim that mixed models take into account correlations within their repeated data. The specification of the structure of the covariance matrix allows to control only partly for that. I notice that the authors did not specify the structure of that matrix: the article they refer to justify the appropriateness of their analyses is not adapted. The specification of the structure of the covariance matrix needs to address, in a mixed model, the difference in handling 4 repeated data per participants that cannot be paired as compared to 4 repeated data that can be paired (two per session with one before and one after the neutral or negative priming sessions, if I count right). Of note is that a covariance structure that is left free of constraint for the fit of the model does not capture appropriately the pairing of the data: it has all chances to capture the covariance in a different way. And a covariance structure that has constraints has more chances to lead to a model that cannot be estimated because of an absence of convergence of the algorithms.

By the way, a single two-sample t-test (or a Mann-Whitney test if appropriate), and not a set of multiple paired-sample t-test as the authors suggest, would answer the goal of the authors to test for what they call the three-way interaction in their comment. This test would be performed between the two groups of participants (BN/controls) with the computation for each participant separately: (assessment after neutral induction-assessment before neutral induction)-(assessment after negative induction-assessment before negative induction). This analysis answers points 1, 2 and 4 they raise together with my point of controlling for the paired data. I would have agreed with their choice of a mixed model if they had an unbalanced dataset within each participant.

We thank the Reviewer for this clarification, and we apologize that our previous response did not adequately distinguish between two different sets of analyses: (1) analyses of DDM parameter estimates, which involved four observations per participant (2 affect conditions × 2 food types); (2) trial-level analyses of choice and response time behavior, where each participant contributed many trials per condition and the dataset is genuinely unbalanced across participants due to trial exclusions – precisely the situation where mixed-effects models with participant-level random slopes are appropriate. The concern about covariance structure applies specifically to the DDM parameter analyses, but does not apply to our trial-level analyses.

We also want to clarify a point about the task design that may have caused confusion. The Food Choice Task was administered only once per session, after the mood induction (i.e., once after negative mood induction, and once after neutral mood induction). As detailed in Figure 1, the task was not completed pre-induction. The four observations per participant in the DDM parameter analyses therefore reflect 2 affect conditions × 2 food types assessed within each condition, not a pre/post structure. This does not change how we address the concern about covariance structure, as there is still a nested feature of interest (food type within condition), but we wanted to correct this misunderstanding explicitly.

For the DDM parameter analyses, we agree with the Reviewer that the original random effects structure did not adequately account for the paired nature of the four within-person observations.

We have addressed this in two ways.

First, we re-estimated the mixed model specifying an unstructured covariance matrix using the nlme package, which places no constraints on the correlation pattern among the four withinperson observations. We acknowledge the Reviewer's point that an unconstrained covariance matrix is not guaranteed to recover the within-session pairing structure. We explored whether a more constrained specification would be preferable. Specifically, we tested a nested random effect of affect condition within subject, which would directly encode the pairing of Low-Fat and High-Fat observations within each session. However, this model failed to converge. This is not a numerical issue but a fundamental identification problem: with only two observations per session per subject, the session-level and residual variance components cannot be separately estimated. We therefore selected the unstructured model as a more conservative option. Importantly, even if the unstructured model does not explicitly encode the pairing, it is a more general mathematical formula which would not impose incorrect constraints on the correlation structure.

Consistent with our original findings, the mixed model with an unstructured covariance matrix yielded a significant three-way interaction (Group × Condition × Food Type: β = 0.28, SE = 0.12, t = 2.36, p = 0.020). All simple effects analyses have been updated to reflect the models with this covariance structure, and these are reported in the updated Supplementary Tables.

Second, following the Reviewer's suggestion (adapted to the actual design structure, in which the Food Choice Task was administered once per session after the mood induction rather than before and after), we computed a difference-in-difference score for each participant's relative attribute onset parameter (τs) following the affect inductions: (negative condition, high-fat − negative condition, low-fat) − (neutral condition, high-fat − neutral condition, low-fat). This score directly encodes the paired structure by construction, bypassing the covariance specification problem entirely. Consistent with the Reviewer's recommendation to use a non-parametric test where appropriate, we used a Wilcoxon rank-sum test (equivalent to Mann-Whitney U) to compare these difference scores between groups. The results confirmed that BN participants showed significantly larger food-type-specific changes in τs following negative affect induction relative to HC (W = 156, p = 0.018). We then applied this approach to all other DDM parameters (i.e., ωtaste, ωhealth, α, τND, and z), and report these results alongside updated mixed-effects model results in the Supplementary Materials. The conclusions drawn from the difference-in-difference analyses were consistent with those from the mixed-effects models across all parameters.

Both approaches converge on the same conclusion and we report both sets of complementary results in the manuscript: the updated mixed-effects models address the full factorial design in a single framework, while the added difference-in-difference analyses explicitly resolve the covariance specification problem by encoding the paired structure directly into each participant’s score, as the Reviewer recommended.

Reviewer #2 (Public review):

Summary:

Binge eating is often preceded by heightened negative affect, but the specific processes underlying this link are not well-understood. The purpose of this manuscript was to examine whether affect state (neutral or negative mood) impacts food choice decision-making processes that may increase likelihood of binge eating in individuals with bulimia nervosa (BN). The researchers used a randomized crossover design in women with BN (n=25) and controls (n=21), in which participants underwent a negative or neutral mood induction prior to completing a food-choice task. The researchers found that despite no differences in food choices in the negative and neutral conditions, women with BN demonstrated a stronger bias toward considering the 'tastiness' before the 'healthiness' of the food after the negative mood induction.

Strengths:

The topic is important and clinically relevant and methods are sound. The use of computational modeling to understand nuances in decision-making processes and how that might relate to eating disorder symptom severity is a strength of the study.

Weaknesses:

Sample size was relatively small, and participants were all women with BN, which limits generalizability of findings to the larger population of individuals who engage in binge eating. It is likely that the negative affect manipulation was weak and may not have been potent enough to change behavior. These limitations are adequately noted in the discussion.

We thank the reviewer for their thorough description of the strengths and weaknesses of this study.

Changes from the earlier version to the revised version

Abstract

Earlier version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and under-consumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods was associated with more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Revised version: Individuals often consume tasty, calorically dense foods in response to negative emotions, a phenomenon exemplified by notions of “stress eating” and “comfort food.” While this link between food and mood can become pathological in binge eating, the decision-making processes underlying this link are poorly understood. Here, we investigated the impact of acute increases in negative affect on when and how strongly the perceived tastiness and healthiness of foods influence food choices in healthy adults and individuals with bulimia nervosa (BN), an eating disorder characterized by cycles of over- and underconsumption of food. In a randomized crossover design, 25 women with BN and 21 healthy controls completed two sessions where they received either a neutral or negative affect induction and then completed a food choice task. Using a time-varying diffusion decision model, we assessed how negative affect influences food choice dynamics for high- and low-fat foods. In the neutral affect condition, individuals with BN considered tastiness relative to healthiness of high-fat foods sooner than healthy controls but maintained a restrictive food choice policy by reducing the weight on tastiness. After a negative affect induction, both groups showed a stronger bias towards considering tastiness before healthiness, but this bias was exaggerated in individuals with BN. This affect-induced bias for high-fat foods predicted more frequent subjective binge episodes over three months. These results provide insights into how negative emotion influences food choices and may explain why binge eating in BN is more likely during high negative affect, while dietary restriction is more likely during low negative affect.

Introduction

Every day, multiple times a day, people must make decisions about what to eat. A growing body of research suggests that these decisions involve considerations of different attributes of the available food options (e.g., healthiness and tastiness), and that these attributes often present the decision-maker with challenging conflicts (e.g. “Should I choose the healthier, less delicious food or the better tasting, less healthy food?”). Difficulties resolving these conflicts can take on pathological forms in individuals with eating disorders, and some eating disorders are characterized by food choices that are biased toward one extreme (e.g., toward subjectively less tasty, but healthier, low-fat foods in anorexia nervosa; (Foerde et al., 2015)). However, food choices seem to episodically oscillate between extremes in the case of bulimia nervosa (BN), which is commonly marked by prolonged periods of rigid dietary restriction punctuated by out-of-control binge eating of highly palatable foods. Although subclinical forms of this oscillation are ubiquitous in examples of “yo-yo dieting” and “dietary lapses,” surprisingly little cognitive neuroscience research to date has attempted to explain it.

Self-report data suggest that rising negative affect is one reliable predictor of the switch from more to less restrictive food choices (Konttinen et al., 2010). Although some individuals reduce intake in response to stress (Hill et al., 2022), research using both animal models and healthy adult samples show that stress can precipitate the overconsumption of highly palatable foods (Dionysopoulou et al., 2021; Jacques et al., 2019; Tomiyama, 2019). This connection between negative emotion and food choice is even more pronounced in binge eating. Inducing negative affect in the laboratory environment disproportionately increases food intake among individuals with binge eating (Telch & Agras, 1996). Ecological momentary assessment data show that binge eating in the natural environment is also much more likely to occur in the context of increasing and unstable negative emotions (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). In addition, reactivity to negative emotions, or negative urgency, has been shown to predict or correlate with more frequent binge eating in individuals with BN (Fischer et al., 2013; Schnepper et al., 2021).

Although momentary increases in negative emotions are closely tied to the overconsumption of highly palatable foods, the cognitive processes underlying this tight link, particularly in the case of binge eating, are not well understood. Studies to date have failed to find support for the notion that increased negative affect, specifically stress, impairs inhibitory control processes in healthy adults (Allen et al., 2022) or in BN (Dreyfuss et al., 2017; Westwater et al., 2021). While another study found that healthy adults and individuals with BN differed in how stress impacted visual processing-related neural activity, they failed to find support for the hypothesis that increased negative affect abnormally increased limbic neural responses to palatable food stimuli in BN (Collins et al., 2017).

Earlier version: We recently tested another possibility: that negative affect abnormally impacts decision-making in BN, specifically promoting more high-fat, tasty food choices that are typically consumed during binge-eating episodes (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Revised version: We recently tested the possibility that negative affect abnormally impacts decision-making in BN by promoting food choices that include the high-fat, appetitive foods that are typically consumed during binge-eating episodes (Gianini et al., 2019). Individuals with BN and healthy controls completed a food choice task after a negative and neutral affect induction (Figure 1). Counter to expectations, negative affect did not influence ultimate food choices in either group, and women with BN showed restrictive food choices whether neutral or negative affect was induced: they chose low-fat foods more often (Figure 2A), and these choices were strongly predicted by healthiness ratings (Figure 2B). However, these analyses focused primarily on individuals’ ultimate choices and may have overlooked the possibility that even small changes in negative affect could influence more subtle aspects of the food decision-making process leading up to those choices.

Affect induction and task design.

(A) Study timeline. BN and HC participants completed the study tasks on two separate days. On session 1, participants were randomly assigned to the neutral mood or negative mood induction before completing the Food Choice Task. On session 2, participants experienced the alternative mood induction before completing another run of the Food Choice Task. The mood inductions involved combinations of music and autobiographical writing. (B) Food Choice Task. During the Food Choice Task, both BN and HC participants rated 43 food items across three phases. In the Tastiness Ratings and Healthiness phases, participants rated each item on a 5-point Likert scale from Bad to Good and Unhealthy to Healthy, respectively. In the Choice phase, participants indicated their strength of preference for a presented food item, compared to a personally tailored neutral reference item.

Earlier version: (A) Food choice task behavior estimated from regression models. Re-analysis of the raw data excluding outlier response times (2.5% of trials n = 85) replicated original findings (Gianini et al., 2019). While both groups were less likely to choose high-fat foods (over the neutral reference item) than low-fat foods (over the neutral reference item), the BN group was even less likely than the HC group to choose high-fat food items. However, we did not identify any significant effects of the Affect Condition on choices. Error bars indicate 95% confidence intervals of the estimated effects. (B) Influence of health and taste ratings on food choice. Health ratings influenced food choice more in the BN group than in the HC group. Within the HC group, food choice was influenced more strongly by taste ratings than health ratings. Error bars indicate standard errors of the estimated coefficients. Note. Corresponding statistics are presented in Table S8. HC = healthy controls; BN = bulimia nervosa.

Revised version: (A) Food choice task behavior estimated from regression models. Re-analysis of the raw data excluding outlier response times (2.5% of trials n = 85) replicated original findings (Gianini et al., 2019). While both groups were less likely to choose high-fat foods (over the neutral reference item) than low-fat foods (over the neutral reference item), the BN group was even less likely than the HC group to choose high-fat food items. However, we did not identify any significant effects of the Affect Condition on choices. Error bars indicate 95% confidence intervals of the estimated effects. (B) Influence of health and taste ratings on food choice. Health ratings influenced food choice more in the BN group than in the HC group. Within the HC group, food choice was influenced more strongly by taste ratings than health ratings. Error bars indicate standard errors of the estimated coefficients. Note. Corresponding statistics are presented in Table S9. HC = healthy controls; BN = bulimia nervosa.

Earlier version: These more subtle features can be quantified with cognitive computational modeling. Sequential sampling models leverage both choice and response time data to estimate parameters that capture individual differences in decision-making dynamics (Forstmann et al., 2016; Ratcliff & Smith, 2004). These models assume that decision-making involves the noisy accumulation of evidence over time until a predetermined threshold is reached, leading to a response. Sequential sampling models assume that both ultimate responses and response times are meaningful signals about the cognitive processes underlying decision-making. Though originally developed to describe perceptual, psychomotor, and memory tasks, these models can be applied to value-based decision-making tasks by incorporating attribute ratings for the stimuli presented in a trial. This allows researchers to examine how and when different attributes influence the decision-making process (Busemeyer et al., 2019; Clithero, 2018).

Revised version: These more subtle features can be quantified with computational cognitive modeling. Sequential sampling models leverage both choice and response time data to estimate parameters that capture individual differences in decision-making dynamics (Forstmann et al., 2016; Ratcliff & Smith, 2004). These models assume that decision-making involves the noisy accumulation of evidence over time until a predetermined threshold is reached, leading to a response. Therefore, both ultimate responses and response times are meaningful signals about the cognitive processes underlying decision-making. Though originally developed to describe perceptual, psychomotor, and memory tasks, these models can be applied to valuebased decision-making tasks by incorporating participants’ ratings of different attributes for the stimuli presented in a trial. This allows researchers to examine how and when these attributes influence the decision-making process (Busemeyer et al., 2019; Clithero, 2018).

Research using sequential sampling models indicates that food choices are influenced not only by the weights assigned to attributes like healthiness and tastiness, but also by the time at which these attributes enter the decision-making process (Chen et al., 2022; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Maier et al., 2020; Sullivan et al., 2015; Sullivan & Huettel, 2021). Specifically, using time-varying sequential sampling models, studies have demonstrated that healthy individuals assign different subjective weights to healthiness and taste attribute ratings of foods (i.e., they have differing degrees to which they influence the evidence accumulation rate); that healthiness and taste attribute information can enter the decision process at different times; and that these relative weighting strengths and timing of attribute consideration have distinct influences on ultimate food choices (Maier et al., 2020). This notion is consistent with theories of both emotion- and food-craving regulation, which posit that there are separate attention deployment and stimulus valuation steps involved in these regulation processes (e.g., (Giuliani & Berkman, 2015; Gross, 2015; Han et al., 2018).

Notably, even if one of these decision parameters (e.g., weights on healthiness or tastiness, attribute onsets) biased individuals toward selecting a low-fat (or high-fat) food, the other parameter could compensate for it to ultimately generate the opposite response (Figure 3). For example, if taste information was considered before healthiness information, but healthiness information was weighted strongly enough, a decision-maker could choose a healthier, instead of a tastier, food (Figure 3B). These latent aspects of the decision-making process, either orthogonally or in parallel, could be sensitive to state changes like increases in negative affect and could be altered in BN. Specifically, these latent decision parameters could help explain why dysregulated consumption of tasty, high-fat foods in BN is more likely during periods of high negative affect, but restricted intake of lower-fat foods is more likely during periods of low negative affect. In one case, increasing negative affect could alter an individual’s attribute weights, either reducing their weight on healthiness (dashed blue line) or increasing their weight on tastiness (dashed red line), which would increase their likelihood of choosing binge-type foods that are considered to be highly palatable yet also unhealthy (Figure 3C). Alternatively, increasing negative affect could cause individuals with BN to consider healthiness information later in the decision-making process, giving taste information more time to bias choices towards binge-type foods (Figure 3D). Clarifying these mechanisms is critical for translational efforts because data suggest that the weighting of attributes and the time at which they are considered have different neural substrates (Maier et al., 2020), and the treatment offered to patients should depend on which scenario more accurately reflects their pathology.

Example evidence accumulation trajectories predicted by the starting time diffusion decision model (stDDM).

Each trajectory is a simulated agent that considers one attribute alone before beginning to consider both tastiness and healthiness attributes together. (A) Tastiness onset delay. The dashed, green trajectory represents a case where positive healthiness information about a low-fat food item quickly biases the agent towards the accept threshold before taste information has time to influence the evidence accumulation process. The solid purple trajectory illustrates a case where evidence related to a high-fat food is initially biased towards the reject threshold while the aversive healthiness information dominates the evidence accumulation process. Once information about the item’s appetitive tastiness comes online, the evidence accumulation changes its trajectory towards the accept threshold. (B) Healthiness onset delay. The dashed, orange trajectory illustrates a case where evidence related to the highly appetitive tastiness attribute dominates the evidence accumulation process, influencing the trajectory to terminate at the accept threshold of a high-fat food before the healthiness attribute is considered. For the solid brown trajectory, a less appetitive high-fat food item is ultimately rejected once aversive healthiness information enters the evidence accumulation process. (C-D) Hypothesized model of affect-induced binge-eating. Each trajectory is a simulated agent with BN making a decision involving a high-fat food. During neutral affect (gray solid line), the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (C) Attribute-weight hypothesis. During negative affect, either the attribute weight for taste information increases (red dashed line) or the attribute weight for healthiness information decreases (blue dashed line). In both cases, the trajectory is ultimately biased towards the accept threshold. (D) Attributeonset hypothesis. During negative affect, the initial delay shifts and taste information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Here, we applied a diffusion decision model (DDM) with a time-varying drift rate to food choice and response time data from our controlled, randomized crossover study in women with BN and healthy controls (Figure 1A). This model allowed us to investigate the potential exaggerated influence of negative affect on the dynamics of food-specific decision-making in BN. Specifically, we tested how food type (low-fat/high-fat) and affective state (neutral/negative) influenced decision-makers’ attribute weights and onsets.

Our hypotheses were derived from the clinical phenomenology underlying BN. As outlined above, individuals with BN seem to vacillate between unstable decision dynamics: many individuals with the disorder maintain a relatively restrictive diet until negative emotional states induce binge-eating episodes and subsequent compensatory behaviors (i.e., purging). Given this instability, we predicted that our BN participants’ food choices would be influenced by components of the decision-making process whose opposing dynamics could be disrupted by increasing negative affect.

We first hypothesized that in states of neutral affect and for high-fat foods, one of these opposing dynamics would compensate for the other to promote ultimate choices that are restrictive. This could be achieved by either a) weighing taste information more strongly but considering it later than healthy controls, or b) weighing tastiness less strongly but considering it earlier than healthy controls. Second, we hypothesized that, in BN, increasing negative affect would abnormally impact at least one of these latent aspects of decision-making to increase the bias towards high-fat foods. Even though prior analyses indicate that high-fat foods were not ultimately chosen more often, we predicted that after a negative affect induction, individuals would either a) be slower to consider healthiness information, but still weight taste information just as strongly, or b) they would put less weight on healthiness information, but still consider taste information sooner.

Because these combinations require one feature of the decision process to compensate for another, they would result in a much less stable dynamic compared to one in which information weights and onset times are concordant. For example, when healthiness information is considered sooner and is more strongly weighted than taste information, foods low in healthiness would almost always be rejected. In contrast, when healthiness information is considered later but is more strongly weighed than taste information, food choice behavior can be more easily perturbed with state-based changes in weights or the delay in attribute onsets. Here, we show evidence for these unstable decision dynamics that can account for a bias in BN toward low-fat, healthier foods in states of relatively low negative affect, but toward high-fat, tastier foods in states of high negative affect. Although negative affect delayed the onset of the consideration of healthiness information among both HC and BN participants, this effect was more pronounced in the BN group. Ultimately, we found that although individuals with BN consistently weighed healthiness information more strongly, negative affect delayed its entry into the evidence accumulation process. With this longer delay, taste information has more time to bias decision-makers towards high-fat food choices before healthiness information can come online.

Results

Negative affect aberrantly delays consideration of healthiness information among women with BN

Earlier version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who participated in a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). Mood was assessed before and after the affect inductions using the Profile of Mood States (POMS; (McNair et al., 1989)). Negative affect scores were calculated by summing the five negative subscale scores (e.g., anger, confusion, depression, fatigue, and tension) and subtracting the positive subscale score (e.g., vigor). Negative affect scores increased significantly for both groups following the negative affect induction (Gianini et al., 2019, reproduced in Table S1).

Revised version: We analyzed data from individuals with BN (n = 25) and healthy controls (HC, n = 21) who completed a computerized Food Choice Task following negative and neutral affect inductions, across two visits (Figure 1). Mood was assessed before and after the affect inductions using the Profile of Mood States (POMS; (McNair et al., 1989)). Negative affect scores were calculated by summing the five negative subscale scores (e.g., anger, confusion, depression, fatigue, and tension) and subtracting the positive subscale score (e.g., vigor). Negative affect scores increased significantly for both groups following the negative affect induction (Gianini et al., 2019, reproduced in Table S1). Critically, post-induction negative affect within the BN group was significantly higher following the negative affect induction than after the neutral affect induction (mean difference = 17.40, SE = 4.21, t = 4.13, p < 0.001, Cohen’s d = 0.83; see Supplementary Materials for full details), confirming that BN participants completed the food decision task under meaningfully distinct affective states across the two sessions.

Earlier version: To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. The stDDM was fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction. We specified a stDDM where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter (τs) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteωHealth⋅VDHealthωTaste⋅VDTaste+ωHealth⋅VDHealthIf τS<0∧0<t<|τS|If τS>0∧0<t<τSIf >|τS| (1)

Revised version: To quantify the specific computations performed during Food Choice Task decision-making, we used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008)). This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process (Chen et al., 2022; Lombardi & Hare, 2021; Maier et al., 2020). Using this starting time DDM (stDDM), we can determine how individuals weigh healthiness versus taste information and when those attributes influence food-related decision-making in BN and HC groups in each affective state. The stDDM was fitted to data that included individual-level response times, choices, z-scored healthiness and tastiness ratings, and affect induction.

Added in the revised version: We specified a stDDM where individual parameters were drawn from four separate hyperparameters that varied both by group (BN or HC) and by condition (neutral or negative affect induction). To test our hypotheses about the interactions between affect and food type (low-fat or high-fat) on food choice, we allowed each attribute weight parameter (ωTaste and ωHealth) and the relative-starting time parameter τs) to vary as a function of food type. The drift rate determining the evidence update can be written as follows: v(t)={ωTaste⋅VDTasteIf τs<0∧0<t<|τs|ωHealth⋅VDHealthIf τs>0∧0<t<τsωTaste⋅VDTaste+ωHealth⋅VDHealthIf t<|τs| (1)

Earlier version: where ωTaste is the subjective weight given to tastiness, ωHealth is the subjective weight given to healthiness, VDTaste and VDHealth are the value differences in taste and healthiness attribute ratings, respectively, and τs is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τs > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τs, whereas τs < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τs|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in our statistical analyses. We ran three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values -1 and 1, allowing us to separately assess group differences across condition. Following significant interactions, we conducted simple effects analyses to examine the effect of each factor at specific levels of the other factors.

Revised version: where ωTaste is the subjective weight given to tastiness, ωHealth is the subjective weight given to healthiness, VDTaste and VDHealth are the value differences in taste and healthiness attribute ratings, respectively, and τs is the time at which the taste and healthiness attributes come into the evidence accumulation process. If τs > 0, healthiness information is accumulated first, and evidence from taste comes into consideration at time τs, whereas τs < 0 means that taste information is accumulated first and evidence from healthiness information starts to come into the decision process at time |τs|. To test our hypotheses, medians of the posterior distributions for subject-level parameters were used as the dependent variables in our statistical analyses. We ran three separate, fully factorial (Group-by-Affect Condition-by Food Type), mixed-effects linear regressions, one for each parameter of interest. Each of these binary categorical variables were contrast coded using the values -1 and 1, allowing us to separately assess group differences across condition.

Results exploring additional parameters (boundary separation (α), non-decision time (τND), and starting point bias (z)) are presented in the Supplementary Materials along with model comparisons, model identifiability, posterior predictive checks, and parameter recovery exercises.

Added in the revised version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. We found a significant three-way interaction between group, condition, and food type (Table S2; β = 0.32, t = 2.23, p = 0.028), indicating that the two groups differed in how negative affect influenced their attribute bias patterns across food types. To confirm the presence of a three-way interaction, we calculated a difference-in-difference score for each participant’s τs parameter: (negative condition, high-fat – negative condition, low-fat) – (neutral condition, high-fat – neutral condition, low-fat). A Wilcoxon rank-sum test comparing these scores between groups confirmed that BN participants showed significantly larger food-type-specific changes in attribute onset timing following negative affect induction relative to HC (W = 156, p = 0.018). Visual inspection of the data suggested this interaction was driven by group differences across food types in the neutral condition that disappeared in the negative affect condition, and by a more pronounced influence of negative affect on attribute biases in BN (Figure 4A). Specifically, in the neutral condition, the HC group showed greater taste bias for low-fat versus high-fat foods (M = -0.456 vs. -0.314; Table S3), while individuals with BN showed greater taste bias for high-fat versus low-fat foods (M = -0.314 vs. -0.462; pairwise difference t132 = -2.33, p = 0.022). The negative affect induction eliminated this crossover pattern, increasing taste biases across both food types and groups, but with greater increases observed in the BN group (Tables S2, S3).

Removed in the revised version: First, we assessed how relative start time parameters varied across groups as a function of food type in each condition. We found a significant Group-by-Food Type interaction (Table S2; β = −0.30, t = −2.97, p = 0.003), indicating that in the neutral affect condition, the BN group showed a stronger bias towards first considering the taste attribute of high-fat foods than low-fat foods (Figure 4A). A simple effects analysis (Table S3) confirmed that, in the neutral affect condition, the BN group had a significantly stronger bias for taste information for high-fat foods (M = 0.16, t = 2.25, p = 0.026), while the HC group had a non-significantly stronger bias for low-fat foods (M = −0.14, t = −1.89, p = 0.061).

Parameter estimates.

Colors indicate diagnosis: purple = bulimia nervosa (BN); green = healthy controls (HC). Error bars represent the standard error of the mean. (A) Attribute onset. In the neutral condition, we observed a Group by Food Type cross-over effect: while the BN group showed a greater initial bias towards accumulating tastiness information of high-fat foods than low-fat foods, the HC group showed a greater tastiness information bias for low-fat foods than high-fat foods. After the negative affect induction, any Food Type-based distinctions disappeared, and both groups’ biases towards tastiness information increased, but this effect was more pronounced in the BN group. (B) Weight on tastiness information. The HC group put more weight on tastiness information than the BN group. The negative affect induction reduced tastiness weights for the HC group, but not the BN group. (C) Weight on healthiness information. The BN group put more weight on healthiness information than the HC group, especially for high-fat foods. The negative affect induction did not have significant effects on healthiness weights for either group.

Removed in the revised version: A significant interaction between Group and Affect Condition (Figure 4A; β = −0.22, t = −2.22,p = 0.028) suggests that while both groups showed a bias towards earlier consideration of taste information after the negative affect induction, this effect was more pronounced in the BN group. Additionally, this negative affect-induced bias favoring taste information significantly reduced the group-dependent Food Type distinctions that were present in the neutral state, as indicated by a significant three-way interaction of Food Type, Group, and Affect Condition (β = 0.32, t = 2.22, p = 0.028). Due to the significant three-way interaction, we examined how the influence of the negative affect induction differed between groups across food types (Table S4). For low-fat foods, the negative affect-induced shift in attribute timing was significantly different between the BN and HC groups (M = −0.22, t = −2.17, p = 0.032). For high-fat foods, the difference in affect-induced shifts between groups was not significant (M = 0.09, t = 0.91, p = 0.367). In other words, while there was a significantly stronger shift towards earlier taste consideration following the negative affect induction for the BN group compared to healthy controls, the magnitude of this group difference was strongest for low-fat foods.

Reduced tastiness weights are abnormally unaffected by negative affect in BN

Earlier version: Our second and third regression models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C, Table S4–7). In the neutral affect condition, simple effects analyses revealed that the BN group placed significantly less weight on taste information compared to the HC group (Table S5; M = 0.41, t = 8.27, p = 4.14 × 10−13), while simultaneously placing greater weight on health information (Table S7; M = −0.51, t = −9.78, p = 2.22 × 10−17). For ωtaste, we observed a significant main effect of Food Type (β = −0.23, t = −3.69, p = 3.26 × 10−4), indicating that both groups attribute less weight to taste information for high-fat foods than low-fat foods. A simple effects analysis confirmed this effect was present in both groups (HC: M = 0.22, t = 4.98, p = 1.78 × 10−6; BN: M = 0.35, t = 8.61, p = 1.15 × 10−1). Notably, only the BN group had negative weights on high-fat taste information. For ωhealth, a significant Group-by-Food Type interaction (β = 0.29, t = 2.94, p = 0.004) showed that the BN group exhibited particularly strong health-information weighting for high-fat foods compared to healthy controls. Follow-up simple effects analyses (Table S6) confirmed that this food-type difference for ωhealth was significant for the BN group (M = −0.26, t = −5.42, p = 2.43 × 10−7) but not the HC group (M = 0.01, t = 0.11, p = 0.915). This reduced weight on tastiness and increased weight on healthiness information in the BN group may explain how, despite focusing for longer on taste information, individuals with BN ultimately made restrictive eating choices in states of low negative affect.

Revised version: Our second and third regression models assessed how each attribute weight varied across groups as a function of food type in each condition (Figure 4B-C). For taste weights, the HC group placed consistently higher weight on taste compared to the BN group (Table S4–S5; β = -0.39, t = -6.86,p < 0.001; all pairwise contrasts p ≤ 0.002), and both groups placed less weight on taste for high-fat versus low-fat foods β = -0.23, t = -4.05, p < 0.001; all pairwise contrasts p < 0.001). A significant interaction between group and condition (β = 0.18, t = 2.46, p = 0.015) reflects divergent responses to negative affect: HC showed a significant reduction in taste weighting following the negative affect induction (t132 = 3.20, p = 0.002), but the BN group did not (t132 = 0.61, p = 0.542). Consequently, the group difference in taste weighting, while significant in both conditions (all pairwise contrasts p < 0.001), was numerically smaller in the negative affect condition (M = 0.21) than in the neutral condition (M = 0.39; interaction: t132= 1.95, p = 0.053). In contrast, healthiness weights were insensitive to negative affect induction across both groups. However, a significant interaction between group and food type (Table S6–S7; β = 0.28, t = 3.09, p = 0.002) indicated that the BN group placed disproportionately higher weights on health information for high-fat relative to low-fat foods (t132 = -4.40, p < 0.001), while the HC group showed no difference in health weighting across food types (t132 = 0.01, p = 0.897).

Removed in the revised version: For ωtaste, a significant interaction between Group and Condition (β = 0.18, t = 2.14, p = 0.036) indicated that the BN and HC groups’ weight on taste information responded differently to the negative affect induction. The three-way interaction with Food Type was not significant (β = −0.08, t = −0.72, p = 0.474), suggesting that this Group-by-Condition interaction was consistent across food types. Simple effects analyses (Table S5) showed that the healthy controls exhibited a significant reduction in ωtaste (M = 0.16, t = 3.62, p = 4.00 × 10−4), while individuals with BN maintained similar weights on taste information across Conditions (M = 0.03, t = 0.66, p = 0.508). Although individuals with BN generally had lower weights for taste information than healthy adults, the negative affect induction only influenced how healthy controls weighted taste information. Neither group exhibited any affect-based changes in ωhealth (Tables S6).

More severe symptoms of BN are linked to affect-induced delays in processing healthiness information

Earlier version: To assess the clinical significance of our findings, we sought to connect stDDM parameters to retrospective self-reported symptom severity. For these exploratory analyses, we used median values of the individual-level posterior distributions of the attribute onset parameter τs, the parameter which showed an exaggerated response to the affect induction in the BN group. These parameter values were then separately added to mixed-effects negative binomial regression models with Group and Food Type to assess their association with two symptom severity measures in the BN group: the frequency of objectively large binge eating episodes (OBEs) and subjectively large binge eating episodes (SBEs) in the past three months. We examined both types of out-of-control eating episodes because they both contribute to a diagnosis of BN according to ICD-11 (World Health Organization, 2022), and some data suggest that SBEs are more strongly related to negative affect (Brownstone et al., 2013; Brownstone & Bardone-Cone, 2021; Fitzsimmons-Craft et al., 2014). Indeed, we found that that greater negative affect-based decreases in τs for high-fat foods relative to low-fat foods were associated with more frequent SBE in the past three months (Table S8; β = −46.39, t = −2.80, p = 0.005). In other words, those who focused for longer on the tastiness of high-fat foods after the negative affect induction reported more SBE episodes in the past three months (Figure 5). Conversely, affect-based changes in τs were unrelated to OBE frequency (Table S8).

Revised version: To assess the clinical significance of our findings, we sought to connect stDDM parameters to retrospective self-reported symptom severity. For these exploratory analyses, we used median values of the individual-level posterior distributions of the attribute onset parameter xs, the parameter which showed an exaggerated response to the affect induction in the BN group. These parameter values were then separately added to mixed-effects negative binomial regression models with Group and Food Type to assess their association with two symptom severity measures in the BN group: the frequency of objectively large binge eating episodes (OBEs) and subjectively large binge eating episodes (SBEs) in the past three months. We examined both types of out-of-control eating episodes because they both contribute to a diagnosis of BN according to ICD-11 (World Health Organization, 2022), and some data suggest that SBEs are more strongly related to negative affect (Brownstone et al., 2013; Brownstone & Bardone-Cone, 2021; Fitzsimmons-Craft et al., 2014). Indeed, we found that that greater negative affect-based decreases in τs for high-fat foods relative to low-fat foods were associated with more frequent SBE in the past three months (Table S8; τ = -46.39, t = -2.80, p = 0.005). In other words, those who focused for longer on the tastiness of high-fat foods after the negative affect induction reported more SBE episodes in the past three months (Figure 5). Conversely, affect-based changes in xs were unrelated to OBE frequency (Table S8).

Earlier version: Affect-induced changes in information onset predict subjective binge episodes.

Revised version: Affect-induced changes in information onset were associated with more frequent subjective binge episodes.

In the negative affect condition (right facet), longer delays in accumulating healthiness information (i.e., reduced τs) for high-fat foods compared to low-fat foods was associated with more frequent subjective binge episodes. Line type and shape refer to Food Type: Solid lines and circles = Low fat foods; dashed lines and triangles = High fat foods.

Discussion

Earlier version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knäuper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Revised version: Negative affect is a known precursor to over-eating and dietary rule-breaking (Frayn & Knäuper, 2018). Healthy adults, and even rodents, can exhibit dysregulated food consumption and altered food choices after periods of stress or other negative emotions (Dionysopoulou et al., 2021; Jacques et al., 2019; Konttinen et al., 2010; Tomiyama, 2019). Negative affect’s influence on eating behavior is particularly pronounced among individuals with clinical diagnoses like BN, who tend to have large and out-of-control eating episodes in states of strong negative emotion (Cardi et al., 2015; Wonderlich et al., 2022). Prior studies have demonstrated that outside of binge-eating episodes, and in states of lower negative affect, individuals with BN typically engage in restrictive eating behaviors, consuming some low-fat foods and largely avoiding high-fat foods (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990). We previously found that women with BN demonstrated these restrictive food choices on a food-related decision task, regardless of their affective state (Gianini et al., 2019). However, by leveraging a computational model that captures the decision-making process underlying ultimate choices, we found that negative affect more strongly biased individuals with BN than healthy adults, increasing the amount of time that taste information alone influenced their decisions involving food. This affect-induced delay in the consideration of healthiness versus tastiness may help account for the fact that negative affect is a common precursor to taste-driven overeating in healthy adults, and the fact that it is a particularly strong predictor of taste-driven, out-of-control eating in individuals with BN. Consistent with this notion, we found that negative affect-induced delays in accumulating health information relative to taste information for higher fat foods were associated with more frequent out-of-control eating episodes in BN.

Earlier version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and negative weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) likely facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste.

Revised version: Previous studies have shown that in states of neutral affect, individuals with BN are more likely to consume foods that are low-fat and considered to be healthy (e.g., salads, fruits), as well as use healthiness information to guide their food choices (Alpers & Tuschen-Caffier, 2004; Bjorlie et al., 2022; Elran-Barak et al., 2015; Kales, 1990; Neveu et al., 2018; Schnepper et al., 2021). Our results demonstrate that in a neutral mood state, despite an initial tendency to focus on taste information before healthiness information, the BN group had stronger weights on healthiness information and lower, and in fact negatively signed, weights on taste, especially for high-fat foods. This negative weight on taste information (i.e., an aversive assignment to the value of taste) may have facilitated the BN group’s restrictive food choice policy compared to healthy controls. Individuals with BN may compensate for their tendency to focus on the tastiness of high-fat foods by negatively reappraising or deemphasizing the relative importance of taste and highly weighting their food choices based on health.

Earlier version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although previous studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-of-control eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes— episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). These results suggest that individuals whose task behavior is more reactive to negative affect tend to be the most symptomatic. However, our correlational design does not allow us to determine whether this reactivity causes the symptoms.

Revised version: During binge-eating episodes, which often follow increases in negative affect, individuals with BN tend to consume large amounts of high-fat foods that were previously avoided (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Collins et al., 2017; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007). Although prior studies examined food consumption associated with negative emotions in BN (Cardi et al., 2015), discrete food choice behaviors had not been investigated, and the influence of negative affect on specific aspects of the decision-making process were unclear. We showed that this affect-related change in eating behavior may occur because negative affect disrupts the delicate balance between how much weight tastiness and healthiness carry in influencing food choices and the relative speed with which they begin to exert their influence (Barakchian et al., 2021; HajiHosseini & Hutcherson, 2021; Hutcherson & Tusche, 2022; Lim et al., 2018; Maier et al., 2020; Schubert et al., 2021; Sullivan et al., 2015; Sullivan & Huettel, 2021). When in states of high negative affect, both healthy adults and individuals with BN may consider food’s appealing taste for longer, biasing them towards choosing tastier, high-fat foods. If this temporal advantage for tastiness is long enough, and a positive weight on tastiness is strong enough, individuals with BN could become more vulnerable to out-of control eating because they could decide to eat tastier or high-fat foods before ever considering their perceived healthiness. This notion is consistent with our finding that greater affect-induced delays in processing healthiness information were associated with more frequent subjective binge episodes— episodes of dysregulated eating that may not include a large amount of food, but like objectively large binge-eating episodes, are dominated by carbohydrates and high-fat foods that are highly palatable (Presseller et al., 2023), and are often spurred by negative affect (Alpers & Tuschen-Caffier, 2001; Berg et al., 2013; Haedt-Matt & Keel, 2011; Hilbert & Tuschen-Caffier, 2007; Smyth et al., 2007).

Earlier version: Our findings reveal new insights into the dynamics of food-related decision-making in both healthy adults and BN and may help to inform new interventions. When attribute weights are inconsistent with a larger goal, decision-makers can use shifts in the timing of attribute consideration to achieve a desired outcome (Amasino et al., 2019; Maier et al., 2020). For example, earlier shifts in the relative onset time of healthiness information could help an individual with stronger weights on tastiness adhere to a diet. Perhaps targeting affect-induced shifts in the timing of attribute consideration could help regulate the over- and under-controlled eating seen in binge-fast cycles. For example, past research in individuals without eating disorders indicates that cueing attention towards the healthiness of foods during a choice task both increased the weight on healthiness information and increased its onset time during decision-making (Barakchian et al., 2021; Maier et al., 2020) but see (Sullivan & Huettel, 2021). Several interventions already focus on modifying attentional biases in individuals with obesity and binge eating (Boutelle et al., 2016; Brockmeyer et al., 2019; Stojek et al., 2018). Our results suggest that such interventions could be targeted even more precisely. One way for individuals to avoid emotion-driven binge eating may be to focus attention away from the taste attributes of tastier, high-fat foods during states of high negative affect. Conversely, during states of low negative affect, increased attention on the positive taste attributes of all foods may be instrumental in motivating individuals with BN to engage in less restrictive eating.

Revised version: Our findings reveal new insights into the dynamics of food-related decision-making in both healthy adults and BN and may help to inform new interventions. When attribute weights are inconsistent with a larger goal, decision-makers can use shifts in the timing of attribute consideration to achieve a desired outcome (Amasino et al., 2019; Maier et al., 2020). For example, earlier shifts in the relative onset time of healthiness information could help an individual with stronger weights on tastiness adhere to a diet. Perhaps targeting affect-induced shifts in the timing of attribute consideration could help regulate the over- and under-controlled eating seen in binge-fast cycles. For example, past research in individuals without eating disorders indicates that cueing attention towards the healthiness of foods during a choice task both increased the weight on healthiness information and increased its onset time during decisionmaking (Barakchian et al., 2021; Maier et al., 2020) but see (Sullivan & Huettel, 2021). Several interventions already focus on modifying attentional biases in individuals with obesity and binge eating (Boutelle et al., 2016; Brockmeyer et al., 2019; Stojek et al., 2018). Our results suggest that such interventions could be targeted even more precisely. One way for individuals to avoid emotion-driven binge eating may be to focus attention away from the taste attributes of tastier, high-fat foods during states of high negative affect. Conversely, during states of low negative affect, increased attention on the positive taste attributes of all foods may be instrumental in motivating individuals with BN to engage in less restrictive eating.

Earlier version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction in BN, choice outcomes were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. Although the Food Choice Task was completed immediately after the affect induction, we do not know whether effects persisted because mood was not repeatedly assessed throughout the task. This and previous studies using task-based designs have only assessed mood ratings immediately following affect inductions (Rouhani et al., 2025; Werthmann et al., 2014), but future studies could include continuous assessments throughout the experiment to better quantify potential group differences in how induced affect changes persist and/or fluctuate. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task. Additionally, future research should include continuous assessments of mood, as well as other relevant factors such as cravings (Konova et al., 2018; Wonderlich et al., 2017), to enrich the connection between mood, discrete food choice, and binge eating.

Revised version: Although we found an aberrant influence of negative affect on the computations that contribute to food choice in BN, in our paradigm, the influence of negative affect was not strong enough to change final choices in either group. Even after the negative affect induction which decreased taste weights for HC and not BN, choice outcomes in BN were ultimately determined by the advantage in weighted evidence for healthiness as opposed to the advantage in relative timing for tastiness. We speculate that this may have been for two reasons. First, the affect induction method may not have been potent enough. In the current study, post-induction increases in negative affect (dHC = 1.34; dBN = 1.19) were smaller than those reported after other induction methods, such as viewing affectively salient images with congruent music (d = 3.33; (Zhang et al., 2014)), and they may not have been powerful enough to effect changes in ultimate choices. Although the Food Choice Task was completed immediately after the affect induction, we do not know how long its effect lasted because mood was not repeatedly assessed throughout the task. This and previous studies using task-based designs have only assessed mood ratings immediately following affect inductions (Rouhani et al., 2025; Werthmann et al., 2014), but future studies could include continuous assessments throughout the experiment to better quantify potential group differences in how induced affect changes persist and/or fluctuate. In addition, our negative-affect inducing music and writing exercise may have created an emotional experience that inadequately represents the negative affect that typically precedes emotional eating or binge eating. Qualitative studies of patients with BN indicate that interpersonal stressors and related negative emotions (e.g., resentment) often precipitate binge episodes (Bohon et al., 2021; Wasson, 2003). Perhaps larger increases in negative affect intensity or exposure to the specific types of negative affect which typically precede binge eating could produce a large enough asynchrony in onset timing of attributes and changes in attribute weights to change individuals’ choices (Figure 6). Second, participants knew they would be offered a snack-sized amount of food from a randomly selected trial to consume at the end of the task. This quantity of food may not have provided participants with access to the type of binge-eating experience that negative affect would otherwise precipitate. Further work is needed to test whether the same processes identified in the current study also drive binge eating outside the lab (Wonderlich et al., 2022). Future adjustments to incentive compatibility, with designs that allow for increased consumption beyond snack-sized portions and/or guaranteed access to a wider-array of high-fat foods could clarify whether opportunities to overconsume highly palatable food would generate different patterns of food choice behavior on the task. Additionally, future research should include continuous assessments of other relevant factors such as cravings (Konova et al., 2018; Wonderlich et al., 2017) to enrich the connection between mood, discrete food choice, and binge eating.

Revised model of affect-induced binge-eating.

Each trajectory is a simulated individual with BN who considers the taste attribute alone (pink shaded sides of the panels) before considering both tastiness and healthiness attributes together (blue shaded sides of the panels). In both panels the solid purple trajectory illustrates a decision involving a low-fat food, where tastiness information has a weak, positive weight and healthiness information has a strong, positive weight. The dashed blue trajectory illustrates a decision involving a high-fat food, where tastiness information has a strong, positive weight and healthiness information has a strong, negative weight. (A) Low Negative affect. During neutral affect, the high-fat food is trending towards the accept threshold until the onset of aversive healthiness information biases the trajectory towards the reject threshold. (B) High Negative affect. During negative affect, the initial delay is shifted and tastiness information is accumulated longer before healthiness information comes online. With a longer delay in the onset of healthiness information, the evidence accumulation for the high-fat food has enough time to reach the accept threshold.

Earlier version: Of note, the modest sample size (n = 46) composed exclusively of female participants may limit the generalizability of our findings. Replication in larger samples that includes male participants is needed. Nonetheless, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Revised version: Of note, the modest sample size (n = 46) composed exclusively of female participants may limit the generalizability of our findings. Replication in larger samples that include male participants is needed. Nonetheless, the study’s repeated-measures design benefits from the added statistical power and sensitivity of within-subjects statistical tests, which reduce error variance by controlling for individual differences (Charness et al., 2012; Judd et al., 2017; Westfall et al., 2014). Moreover, we collected a large quantity of data per subject, which enhances the validity of the findings by providing a more precise estimate of the effects of the experimental manipulation, despite the smaller sample size (Baker et al., 2021; Smith & Little, 2018).

Overall, our results advance our understanding of how changes in affective states may shift individuals with BN away from typically restrictive dietary choices and towards binge eating. Specifically, negative affect can bias individuals with BN to focus on taste information for longer, and the strength of this effect for high-fat foods was associated with symptom severity. Our results add to recent data suggesting that computational modeling can detect subtle alterations in decision-making processes in BN that are sensitive to state changes (Berner et al., 2023), and future research should continue examining the potential connections between negative affect and eating-related decision-making using these analytic approaches.

Materials and Methods

Participants

Data were collected from 25 individuals who met DSM-5 diagnostic criteria for BN and 21 healthy controls, group-matched for age and BMI. All participants were female and aged 18-40 years. The sample consisted of 27 participants who identified as Caucasian (58.7%), 7 as Hispanic (15.2%), 6 as Asian (13.0%), 3 as Black (6.5%), and 3 as mixed race (6.5%). For details on recruitment and exclusions, see (Gianini et al., 2019).

Materials and procedures

Participants completed the tasks over two separate study sessions, each lasting approximately four hours. On each study day, participants were instructed to abstain from consuming food or drinks except for a standardized meal consumed two hours preceding the study session.

Earlier version: Participants completed a Food Choice Task (Steinglass et al., 2015) on two separate days, in two counterbalanced states: after a neutral affect induction, and after a negative affect induction. This task is composed of three phases: Healthiness Rating, Tastiness Rating, and Choice. In each phase, participants viewed the same 43 food items, each categorized as either low-fat (< 30% kcals from fat) or high-fat (> 30% kcals from fat). During the Health Rating phase, participants used a 5-point scale to rate the healthiness of each food item from 1 (Unhealthy) to 3 (Neutral) to 5 (Healthy). During the Taste Rating phase, participants used a 5-point scale to rate the tastiness of each food item from 1 (Bad) to 3 (Neutral) to 5 (Good). The order of the Health Rating and Test Rating phases was counterbalanced. After both phases were completed, one food item that was rated Neutral for both health and taste was selected as a Reference Item for use in the Choice phase. If no item was rated Neutral for healthiness and tastiness, then a food item rated 3 on health and 4 or above on taste was selected following previously established procedures (Steinglass et al., 2015). During the Choice phase, participants made choices between the neutral Reference Item and other food items presented on the computer screen. Participants indicated their choice on a 5-point scale from 1 (No – select the Reference Item) to 3 (Indifference) to 5 (Yes - select the shown food item). Participants were informed that one randomly selected trial would be selected for payout at the end of the experiment and would receive the food item they selected on that trial.

Revised version: Participants completed a Food Choice Task (Steinglass et al., 2015) on two separate days, in two counterbalanced states: after a neutral affect induction, and after a negative affect induction. This task is composed of three phases: Healthiness Rating, Tastiness Rating, and Choice. In each phase, participants viewed the same 43 food items, each categorized as either low-fat (< 30% kcals from fat) or high-fat (> 30% kcals from fat). During the Health Rating phase, participants used a 5-point scale to rate the healthiness of each food item from 1 (Unhealthy) to 3 (Neutral) to 5 (Healthy). During the Taste Rating phase, participants used a 5-point scale to rate the tastiness of each food item from 1 (Bad) to 3 (Neutral) to 5 (Good). The order of the Health Rating and Taste Rating phases was counterbalanced. After both phases were completed, one food item that was rated Neutral for both health and taste was selected as a Reference Item for use in the Choice phase. If no item was rated Neutral for healthiness and tastiness, then a food item rated 3 on health and 4 or above on taste was selected following previously established procedures (Steinglass et al., 2015). During the Choice phase, participants made choices between the neutral Reference Item and other food items presented on the computer screen. Participants indicated their choice on a 5-point scale from 1 (No - select the Reference Item) to 3 (Indifference) to 5 (Yes - select the shown food item). Participants were informed that one randomly selected trial would be selected for payout at the end of the experiment and would receive the food item they selected on that trial.

The experiment included an affect induction that included a combination of music and autobiographical writing (Werthmann et al., 2014). During the negative affect induction, participants were asked to write about a recent negative experience while listening to “Adagio for Strings” by Samuel Barber for 8 minutes. During the neutral affect induction, participants were asked to write about the route they took to get to the study site while listening to “Dancing with the Sun” by Celia Felix for 8 minutes. The 65-item Profile of Mood States scale was administered before and after the affect inductions to assess changes in affect (McNair et al., 1989).

Participants also completed the following measures: Eating Disorder Examination Questionnaire (EDE-Q; (Fairburn & Beglin, 1994)); Difficulties in Emotion Regulation Scale (DERS; (Gratz & Roemer, 2004)); Urgency, Premeditation, Perseveration, Sensation Seeking, and Positive Urgency Behavior ScaleNegative Urgency subscale (UPPS-P Negative Urgency; (Lynam et al., 2006)); Beck Depression Inventory (BDI; (Beck et al., 1961)); State Anxiety Inventory (STAI; (Spielberger et al., 1983)).

For additional information, please see (Gianini et al., 2019).

Data analysis

Removed in the revised version: Data transformations and exclusions

Earlier version: Following the procedure from (Gianini et al., 2019), participants’ responses from the Choice phase of the FCT were binarized from their 5-point scale. Responses marked 1 or 2 (i.e., “no”) were determined to be choices of the reference item, responses marked 4 or 5 (i.e., “yes”) were determined to be choices of the presented food item, and responses marked 3 (i.e., “indifferent”) were omitted (Total n =599; BN n = 292, HC n = 307).

Revised version: Data transformations and exclusions: Following the procedure from (Gianini et al., 2019), participants’ responses from the Choice phase of the FCT were binarized from their 5-point scale. Responses marked 1 or 2 (i.e., “no”) were determined to be choices of the reference item, responses marked 4 or 5 (i.e., “yes”) were determined to be choices of the presented food item, and responses marked 3 (i.e., “indifferent”) were omitted (total n =599; BN n = 292, HC n = 307).

In contrast to the original study, we excluded trials containing outlier responses times (RTs). Outlier response times likely reflect phenomena outside the decision-making processes of interest, including attention lapses or accidental responses (Cousineau & Chartier, 2010; Ratcliff, 1993). We excluded outlier responses using cutoffs of ±3 standard deviations (SD) from the mean (Berger & Kiefer, 2021). For each participant, the mean and standard deviation of their RTs were calculated and trials containing responses larger/smaller than the mean ±3 SDs were excluded. Additionally, we removed trials where the RTs were greater than 10,000 ms or less than 250 ms after the mean ±3 SDs treatment. Using this method, we removed 2.5% of trials (total n =85; BN n = 51, HC n = 34).

Removed in the revised version: Computational modeling

Earlier version: We used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008) to study the dynamics of the decision-making during the Food Choice Task. This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process. This model was previously fitted to food choice data (Lombardi & Hare, 2021; Maier et al., 2020) and lottery choice data (Chen et al., 2022) from healthy adults.

Revised version: Computational modeling: We used a time-varying diffusion decision model (DDM; (Ratcliff & McKoon, 2008) to study the dynamics of the decision-making during the Food Choice Task. This time-varying model introduces a starting time parameter that indicates the delay with which different attributes enter the evidence accumulation process. This model was previously fitted to food choice data (Lombardi & Hare, 2021; Maier et al., 2020) and lottery choice data (Chen et al., 2022) from healthy adults.

Earlier version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters, which improves the reliability of parameters estimated from data with low trial counts per participant (Ratcliff & Childers, 2015; Wiecki et al., 2013). Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHeal) and the relative-starting time parameter (τs) were estimated separately based on food-types (low-fat or high-fat). Using model comparisons and posterior predictive checks, we confirmed that this model best fit our data compared to simpler models (see Supplementary Materials).

Revised version: The starting time DDM (stDDM) was estimated using a hierarchical, Bayesian framework implemented with the R package Rjags, which uses the JAGS MCMC sampling algorithm (Plummer, 2003). Our fitting scripts were adapted from those developed by Hsiang-Yu Chen, Gaia Lombardo, and Todd Hare (Chen et al., 2022). This fitting method simultaneously estimates both group- and individual-level parameters, which improves the reliability of parameters estimated from data with low trial counts per participant (Ratcliff & Childers, 2015; Wiecki et al., 2013). Parameter estimates were fitted to data that included individual-level response times, choices, z-scored health and taste ratings, and Affect Condition. We specified a model where individual parameters were drawn from four separate group-level distributions that varied both by Group (BN or HC) and by Affect Induction (neutral or negative). Within this model, attribute weight parameters (ωTaste and ωHealt) and the relative-starting time parameter (τs) were estimated separately based on food-types (low-fat or high-fat). Using model comparisons and posterior predictive checks, we confirmed that this model best fit our data compared to simpler models (see Supplementary Materials).

Earlier version: Following previously described protocols (Chen et al., 2022), we used group-level priors for attribute weight parameters (ωTaste and ωHealth), the relative-starting time parameter (τs), and the affect induction parameters that were drawn from Gaussian distributions with mean = 0 and standard deviation = 1. These parameter values were then divided by their standard deviations before fitting the model. The priors for boundary separation (α) and non-decision time (τND) were drawn from uniform distributions with ranges 1.0 × 10−4 to 5 and 0 to 10, respectively. The priors for starting point bias (Z) were drawn from a beta distribution where both the shape and scale parameters were set to 2. All individual-level priors were drawn from gamma distributions with shape parameters of 1 and scale parameters of 0.1. Posterior estimates were drawn from three chains, each with 100,000 samples (85,000 discarded as burn-in) and thinning every 10 samples. Convergence among chains was assessed using the Gelman-Rubin statistic (Gelman & Rubin, 1992) using a threshold of 1.1. Model fit was quantified with the Widely Applicable Information Criterion (WAIC) (Watanabe, 2010), which penalizes for model complexity based on the number of parameters to prevent overfitting. For parametric tests involving parameter estimates, we used the medians of the individual-level posterior distributions. See the Supplementary Materials for model information on model comparisons, model identifiability, posterior predictive checks, and parameter recovery.

Revised version: Following previously described protocols (Chen et al., 2022), we used group-level priors for attribute weight parameters (ωTaste and ωHeal), the relative-starting time parameter (τs). and the affect induction parameters that were drawn from Gaussian distributions with mean = 0 and standard deviation = 1. These parameter values were then divided by their standard deviations before fitting the model. The priors for boundary separation (α) and non-decision time τND were drawn from uniform distributions with ranges 1.0 × 10-4 to 5 and 0 to 10, respectively. The priors for starting point bias (z) were drawn from a beta distribution where both the shape and scale parameters were set to 2. All individual-level priors were drawn from gamma distributions with shape parameters of 1 and scale parameters of 0.1. Posterior estimates were drawn from three chains, each with 100,000 samples (85,000 discarded as burn-in) and thinning every 10 samples. Convergence among chains was assessed using the Gelman-Rubin statistic (Gelman & Rubin, 1992) using a threshold of 1.1. Model fit was quantified with the Widely Applicable Information Criterion (WAIC) (Watanabe, 2010), which penalizes for model complexity based on the number of parameters to prevent overfitting. For parametric tests involving parameter estimates, we used the medians of the individual-level posterior distributions. See the Supplementary Materials for model information on model comparisons, model identifiability, posterior predictive checks, and parameter recovery.

Earlier version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. These models were implemented using the lme4 (Bates et al., 2015) and lmertest (Kuznetsova et al., 2013) R packages. Each regression model included Food Type (coded -1/1 for low-fat/high-fat foods), Affect Condition (coded - 1/1 for Neutral/Negative), and Group (coded -1/1 for HC/BN), and their interactions as independent variables. For these models, participant-level random intercepts were included to account for within-subject correlations across repeated measurements. Statistical significance was evaluated using α = 0.050.

Revised version: Multilevel linear regressions were run to evaluate our hypotheses regarding the influence of Food Type, affect induction, BN diagnosis, and their interactions on parameter estimates. Each regression model included Food Type (coded -1/1 for low-fat/high-fat foods), Affect Condition (coded -1/1 for Neutral/Negative), and Group (coded -1/1 for HC/BN), and their interactions as independent variables. For these models, the intercepts were treated as random effects and statistical significance was evaluated using α = 0.050.

Added in the revised version: Because the parameters of interest (i.e., ωs, ωTaste, ωHealth) were estimated separately per condition and food type, each participant contributed exactly four observations to these models (2 affect conditions × 2 food types), creating a fully balanced but structured repeated-measures design. Standard random-intercept-only models implemented in lme4 (Bates et al., 2015) and lmerTest (Kuznetsova et al., 2013) do not explicitly model the residual covariance among these four within-person observations. To address this, models were re-implemented using the nlme package (José C. Pinheiro et al., 2022), which permits explicit specification of the within-subject residual covariance structure.

Added in the revised version: We evaluated three candidate covariance structures. We first considered a nested random effect of affect condition within subject, which would directly encode the pairing of Low-Fat and High-Fat observations within each session, but this model failed to converge due to the limit. With only two observations per session per subject, the session-level and residual variance components cannot be separately estimated. Next, we implemented a compound symmetry structure using the corCompSymm function, which assumes equal correlations among all four within-person observations. Finally, we implemented an unstructured covariance matrix using the corSymm function, which places no constraints on the pattern of correlations among the four within-person observations and represents the most general feasible specification for this design. Model comparisons via likelihood ratio test and information criteria (LRT: Δdf = 5, p = 0.057; AIC (-119.08 vs. -118.34) marginally favored unstructured; BIC (-57.99 vs. -73.33) favored compound symmetry) did not strongly favor either tractable structure, and we selected the unstructured model as the more conservative option, as all constrained covariance structures are nested within the unstructured model and represent testable restrictions on the covariance parameters (Molenberghs & Verbeke, 2000; Pinheiro & Bates, 2000). All models additionally included heterogeneous residual variances across the four condition-by-food-type cells using the varIdent function, allowing each cell to have its own residual standard deviation.

Added in the revised version: Difference-in-differences scores were additionally computed for each participant and parameter to provide a complementary analysis that encodes the within-session pairing directly by construction, bypassing the covariance specification problem entirely. These analyses and their results are described in the Supplementary Materials.

Earlier version: To unpack observed interaction effects, we ran simple effects analyses to separately evaluate the influence of independent variables at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991). We used the R package emmeans to compute estimated marginal means for specified factor combinations (Lenth, 2025), evaluating statistical significance using α = 0.050.

Revised version: To unpack observed interaction effects, we ran exploratory simple effects analyses to separately evaluate the influence of independent variables at specific values of other independent variables (Aiken et al., 1991; Brambor et al., 2006; Jaccard & Turrisi, 2003; Winer et al., 1991). We used the R package emmeans to compute estimated marginal means for specified factor combinations (Lenth, 2025), evaluating statistical significance using α = 0.050.

Earlier version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used zero-inflated negative binomial models to predict the three-month frequency of self-report Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). We implemented these models using the R package glmmTMB (Brooks et al., 2017). For each of these two models, individual-level parameter estimates from each condition and food-type were included as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.

Revised version: Additional fixed-effects regressions were run to assess the relationship between parameter estimates and symptom severity. For these exploratory analyses, we used negative binomial models to test the association between parameter estimates and the three-month frequency of retrospectively reported Objective Binge Episodes (OBE) and Subjective Binge Episodes (SBE). We implemented these models using the R package glmmTMB (Brooks et al., 2017), with a zero-inflated term included for the model predicting SBE. For each of these two models, individual-level parameter estimates from each condition and food-type were included as independent variables. We applied a Bonferroni correction to control for the FWE of these exploratory analyses, using α = 0.025.